University of Zagreb\
Faculty of Electrical Engineering and Computing

## Text Analysis and Retrieval 2021/2022
https://www.fer.unizg.hr/predmet/apt/

------------------------------

### Semantics
### LAB 2


*Version: 1.0*

(c) 2022 Josip Jukić, Jan Šnajder

Submission deadline: **May 4, 2022, 23:59 CET** 

------------------------------

### Instructions

Hello visitor, this lab assignment consists of three parts. Your task boils down to filling out the missing parts of code and evaluating the cells. These parts are indicated by the "YOUR CODE HERE" template.

Each subtask is supplemented by several tests that you can run. Apart from that, there are additional test that will be executed after submition. If your solution is valid and it passes all of the visible tests, there shouldn't be any problems with the additional tests.

**IMPORTANT: Don't change the names of the predefined methods or random seeds**, because the tests won't execute properly.

You're required to do this assignment **on your own**.

If you stumble upon problems, please refer to josip.jukic@fer.hr for office hours.

## Tasks

### 1. Paraphrase identification

The goal of paraphrase identification is to determine whether a pair of sentences have the same meaning. In this assignment, we will use the [MRPC](https://paperswithcode.com/dataset/mrpc) dataset, a part of the [GLUE benchmark](https://gluebenchmark.com/).

Load the data frames (train & test) and explore their structure. The column `label` indicates whether a given pair of sentences is semantically equivalent (1) or not (0).

```python
import numpy as np
import pandas as pd


# Load CSV files.
df_train = pd.read_csv("train.csv")
df_test = pd.read_csv("test.csv")

df_train.head()
```

The library that will do the heavy lifting for this task is [Podium](https://takelab.fer.hr/podium/), TakeLab's data loading and preprocessing tool for NLP. We advise you to take a look at the official documenation as well as the [walkthrough](https://takelab.fer.hr/podium/walkthrough.html) examples. Additionally, go through the main primitives: [Vocab](https://takelab.fer.hr/podium/package_reference/vocab_and_fields.html#vocab), [Field](https://takelab.fer.hr/podium/package_reference/vocab_and_fields.html#field), and [Dataset](https://takelab.fer.hr/podium/package_reference/datasets.html).

```python
from podium import Vocab, Field, LabelField
from podium.datasets import TabularDataset
from podium.vectorizers import GloVe
```

#### (a)

Construct three fields: `sentence1` (Field), `sentence2` (Field), and `label` (LabelField). Transform the sentences to lower case using the attribute `pretokenize_hooks`. Refer to the [documentation](https://takelab.fer.hr/podium/preprocessing.html#hooks) to see how this can be achieved.

In order to model semantics, we will utilize [GloVe](https://nlp.stanford.edu/projects/glove/) distributional word vectors.

<!-- grade_id: cell-dff7bf4dffadbf13, points: 0, grade: false, solution: true -->
```python


max_vocab_size = 10_000
vocab = Vocab(max_size=max_vocab_size, min_freq=2)

S1 = None # Sentence1 field
S2 = None # Sentence2 field
LABEL = None # Label field

# YOUR CODE HERE
raise NotImplementedError()


fields = [
    S1,
    S2,
    LABEL,
]

train = TabularDataset.from_pandas(df_train, fields)
test = TabularDataset.from_pandas(df_test, fields)
train.finalize_fields()

glove = GloVe()
# Load only the vectors of vocab words.
embeddings = glove.load_vocab(vocab)

# Generate padded batch.
train_batch = train.batch(add_padding=True)
test_batch = test.batch(add_padding=True)
```

<!-- grade_id: cell-d078098275346006, points: 2, grade: true, solution: false -->
```python
assert "washington" in vocab.stoi
```

Generated batch stores values for the corresponding fields, which can be accessed with the dot notation. See the example below.

```python
train_batch.sentence1, train_batch.sentence2, train_batch.label
```

#### (b)

Implement `cosine_similarity`, which computes cosine similarities of two 2D arrays across their second axis (index = 1). Refer to the tests below to see a concrete example. [np.einsum](https://numpy.org/doc/stable/reference/generated/numpy.einsum.html) and [np.linalg.norm](https://numpy.org/doc/stable/reference/generated/numpy.linalg.norm.html) may be useful for this task. However, you are not obliged to use those methods.

<!-- grade_id: cell-f0747bf03a37381a, points: 0, grade: false, solution: true -->
```python
def cosine_similarity(a, b):
    """
    Receives two 2D numpy arrays and calculates cosine similarity across the second axis.
    For examples, if `a` and `b` have shape (32, 10), the resulting array should have shape (32,).
    
    Returns:
        1D numpy array with cosine similarities
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6710738e2774f997, points: 2, grade: true, solution: false -->
```python
a = np.array(
    [
        [0.79674043, -0.21774995, 0.24626283, -1.92507862, -0.44471655],
        [-0.83365243, -1.05258529, -0.69114343, 1.94794818, 0.81859483],
        [-1.1742791, 0.39978046, -1.14265924, 1.50492221, 0.99339915],
        [0.58896543, 0.8214453, -0.27131406, 0.45817815, -0.21055904],
    ]
)

b = np.array(
    [
        [0.20286607, 0.34114231, -1.14127489, 0.11783557, -0.43729267],
        [0.34177672, -1.66142734, 1.13159559, 0.07148497, 0.24896589],
        [-0.10376178, 0.30639966, 0.54675361, -0.04626362, 0.1408809],
        [-0.6056932, -1.24619744, -0.2720515, 1.26427211, 1.47021337],
    ]
)

sol = np.array([-0.08127636, 0.1919786, -0.19251815, -0.37211115])

assert np.isclose(cosine_similarity(a, b), sol).all()

```

#### (c)
Implement `top_n`, which retrieves the sorted indices of top `n` values in a numpy array. Refer to the test below.

<!-- grade_id: cell-aa7506bc0511eecf, points: 0, grade: false, solution: true -->
```python
def top_n(sims, n=10):
    """
    Receives a numpy array `sims` and finds the indices of the top `n` highest similarities.
    The indices are returned in the ascending order (from lowest to highest index).
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-061d712459e09a1a, points: 1, grade: true, solution: false -->
```python
assert (top_n(np.array([0.1, 0.5, 0.2, 0.3, 0.9]), 2) == np.array([1, 4])).all()
```

#### (d)
We need to somehow transform the sequences to fixed-size vectors. We will explore several simple approaches.

Extract word embeddings for sentence1 and sentence2 (both for train and test), using `train_batch`, `test_batch`, and `embeddings`. Additionally, compute the mean embedding for both fields and both sets. Store them to `sentence1_train_mean`, `sentence2_train_mean`, `sentence1_test_mean`, and `sentence2_test_mean`.

Once you compute the means, retrieve and print out the top 10 most similar sentence pairs from the train set using the previously implemented methods `cosine_similarity` and `top_n`.

<!-- grade_id: cell-63d484fa72557c59, points: 0, grade: false, solution: true -->
```python
sentence1_train, sentence2_train = None, None
sentence1_test, sentence2_test = None, None
sentence1_train_mean, sentence2_train_mean = None, None
sentence1_test_mean, sentence2_test_mean = None, None


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c1f16db9cca4da7a, points: 1, grade: true, solution: false -->
```python
assert sentence1_train_mean.shape == (len(train), 300)
```

#### (e)
Now use the computed means to create representations. You should make two variants: `*_mul`**multiplies the values element-wise** and `*_cat` simply **concatenates** the means from sentence1 and sentence2.
Load the train and test labels from their corresponding batches. Make sure that the label array is 1D.

<!-- grade_id: cell-818c0bfd3a022b26, points: 0, grade: false, solution: true -->
```python
X_train_mul, X_test_mul = None, None
X_train_cat, X_test_cat = None, None
y_train, y_test = None, None


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bb7c231933127781, points: 2, grade: true, solution: false -->
```python
assert X_train_mul.shape == (len(train), 300)
assert X_train_cat.shape == (len(train), 600)
assert y_train.shape == (len(train),)
```

#### (f)
Finally, let's exploit the fruits of our labor by building a classifier that will identify paraphrases. Implement `train_model` that trains an [`sklearn`](https://scikit-learn.org/stable/) model and returns the fitted model. You may use any `sklearn` model, but you need to achieve **binary $F_1$ generalization score higher than 0.4** for both represenations (multiplied and concatenated) to pass the tests. To calculate $F_1$ with `sklearn.metrics.f1_score`, set `average="binary"`. See the example bellow. For additional information, we advise you to use `classification_report`. Remember that the MRPC dataset is quite challenging and we're using a simple approach, so don't be surprised with low binary $F_1$ scores.

<!-- grade_id: cell-1a247f119b6b096f, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import f1_score
from sklearn.metrics import classification_report


def train_model(X_train, y_train):
    """
    Fit and return the fitted sklearn model.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-adfee54218548429, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-440e012cfbed3809, points: 1, grade: true, solution: false -->
```python

```